Make NumPy available:

```python
import numpy as np
```

## Exercise 07.1 (indexing and timing)

Create two very long NumPy 1D-arrays (vectors) `x` and `y` and sum the arrays using:

1. The NumPy addition syntax, `z = x + y`; and
2. A `for` loop that computes the sum entry-by-entry

Compare the time required for the two approaches for arrays of different lengths (use a very long array for 
the timing). The values of the array entries are not important for this test. Use `%time` to report the time.

*Hint:* To loop over an array using indices, try a construction like:

```python
x = np.ones(10)
y = np.ones(len(x))
for i in range(len(x)):
    print(x[i]*y[i])
```

Output:
```
1.0
1.0
1.0
1.0
1.0
1.0
1.0
1.0
1.0
1.0
```

#### (1) Add two arrays using built-in addition operator:

<!-- grade_id: cell-2c856e54f7c3340e, points: 0, grade: false, solution: true -->
```python
x = np.ones(1_000_000)
y = np.ones(1_000_000)
%time z_builtin = x + y
```

Output:
```
CPU times: user 1.95 ms, sys: 4.66 ms, total: 6.61 ms
Wall time: 6.62 ms
```

#### (2) Add two arrays using own implementation:

<!-- grade_id: cell-f34614f9f0068cc4, points: 0, grade: false, solution: true -->
```python
z_loop = np.zeros(len(x))
%time for i in range(len(x)): z_loop[i] = x[i] + y[i]
```

Output:
```
CPU times: user 179 ms, sys: 944 μs, total: 180 ms
Wall time: 180 ms
```

For very large arrays, using built-in addition is much faster than doing it in a 
loop

### Optional extension: just-in-time (JIT) compilation

You will see a large difference in the time required between your NumPy and 'plain' Python implementations. This is due to Python being an *interpreted* language as opposed to a *compiled* language. A way to speed up plain Python implementions is to convert the interpreted Python code into compiled code. A tool for doing this is [Numba](https://numba.pydata.org/).

Below is an example using Numba and JIT to accelerate a computation:

```python
%pip -q install numba 
import numba
import math

def compute_sine_native(x):
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z

@numba.jit
def compute_sine_jit(x):
    z = np.zeros(len(x))
    for i in range(len(z)):
        z[i] = math.sin(x[i])
    return z
    
x = np.ones(10000000)
%time z = compute_sine_native(x)
compute_sine_jit(x)
%time z = compute_sine_jit(x)
```

**Task:** Test if Numba can be used to accelerate your implementation that uses indexing to sum two arrays, and by how much.

```python
...
```

## Exercise 07.2 (member functions and slicing)

Anonymised scores (out of 60) for an examination are stored in a NumPy array. Write:

1. A function that takes a NumPy array of the raw scores and returns the scores as percentages, sorted from 
   lowest to highest (try using `scores.sort()`, where `scores` is a NumPy array holding the scores).
1. A function that returns the maximum, minimum and mean of the raw scores as a dictionary with the 
   keys '`min`', '`max`' and '`mean`'. Use the NumPy array functions `min()`, `max()` and `mean()` to do the 
   computation, e.g. `max = scores.max()`.  
   
   Design your function for the min, max and mean to optionally exclude the highest and lowest scores from the 
   computation of the min, max and mean. 
   
   *Hint:* sort the array of scores and use array slicing to exclude
   the first and the last entries.

Use the scores 
```python
scores = np.array([58.0, 35.0, 24.0, 42, 7.8])
```
to test your functions.

<!-- grade_id: cell-169ebae60810c6be, points: 0, grade: false, solution: true -->
```python
## Exercise 07.2 (member functions and slicing)

def to_percentage_and_sort(scores):
    percentages = (scores / 60.0) * 100
    percentages.sort()
    return percentages

def statistics(scores, exclude=False):
    sorted_scores = np.sort(scores)
    if exclude and len(sorted_scores) > 2:
        sorted_scores = sorted_scores[1:-1]
    return {
        "min": sorted_scores.min(),
        "max": sorted_scores.max(),
        "mean": sorted_scores.mean()
    }
```

<!-- grade_id: cell-af0b6fd8a3cadb1a, points: 0, grade: true, solution: false -->
```python
## tests ##

scores = np.array([35.0, 58.0, 24.0, 42, 7.8])
assert np.isclose(to_percentage_and_sort(scores), [ 13.0, 40.0, 58.33333333,  70.0, 96.66666667]).all()

s0 = statistics(scores)
assert np.isclose(s0["min"], 7.8)
assert np.isclose(s0["mean"], 33.36)
assert np.isclose(s0["max"], 58.0)

s1 = statistics(scores, True)
assert np.isclose(s1["min"], 24.0)
assert np.isclose(s1["mean"], 33.666666666666666667)
assert np.isclose(s1["max"], 42.0)
```

## Exercise 07.3 (slicing)

For the two-dimensional array

```python
import numpy as np
```

```python
A = np.array([[4.0, 7.0, -2.43, 67.1],
             [-4.0, 64.0, 54.7, -3.33],
             [2.43, 23.2, 3.64, 4.11],
             [1.2, 2.5, -113.2, 323.22]])
print(A)
```

Output:
```
[[   4.      7.     -2.43   67.1 ]
 [  -4.     64.     54.7    -3.33]
 [   2.43   23.2     3.64    4.11]
 [   1.2     2.5  -113.2   323.22]]
```

use array slicing for the below operations, printing the results to the screen to check. Try to use array slicing such that your code would still work if the dimensions of `A` were enlarged.

**Check your results carefully against hand computations.**

#### 1. Extract the third column as a 1D array

<!-- grade_id: cell-dfee6b0ed9343682, points: 0, grade: false, solution: true -->
```python
third_column = A[:, 2]
print(third_column)
```

Output:
```
[  -2.43   54.7     3.64 -113.2 ]
```

#### 2. Extract the first two rows as a 2D sub-array

<!-- grade_id: cell-7bf2f9a8c67029f8, points: 0, grade: true, solution: true -->
```python
first_two_rows = A[:2, :]
print(first_two_rows)
```

#### 3.  Extract the bottom-right $2 \times 2$ block as a 2D sub-array

<!-- grade_id: cell-5206fea47d246222, points: 0, grade: false, solution: true -->
```python
bottom_right_block = A[-2:, -2:]
print(bottom_right_block)
```

Output:
```
[[   3.64    4.11]
 [-113.2   323.22]]
```

#### 4. Sum the last column

<!-- grade_id: cell-34f74988e47b9f87, points: 0, grade: false, solution: true -->
```python
last_column_sum = A[:, -1].sum()
print(last_column_sum)
```

Output:
```
391.1
```

#### Compute transpose

Compute the transpose of `A` (search online to find the function/syntax to do this).

<!-- grade_id: cell-dd2999d2da8070f6, points: 0, grade: false, solution: true -->
```python
A_transpose = A.T
print(A_transpose)
```

Output:
```
[[   4.     -4.      2.43    1.2 ]
 [   7.     64.     23.2     2.5 ]
 [  -2.43   54.7     3.64 -113.2 ]
 [  67.1    -3.33    4.11  323.22]]
```

## Exercise 07.4 (optional extension)

In a previous exercise (04.5) you implemented the bisection algorithm to find approximate roots of a mathematical function. Use the SciPy bisection function `optimize.bisect` (https://docs.scipy.org/doc/scipy/reference/generated/scipy.optimize.bisect.html) to find roots of the mathematical function that was used in the previous exercise. Compare the results computed by SciPy and your program from the earlier exercise, and compare the computational time (using `%time`).

<!-- grade_id: cell-b6580accbcb3c2da, points: 0, grade: false, solution: true -->
```python
from scipy import optimize

...
```